Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Python for Data Science Bootcamp
## Assignment 2

NOTE: This assignment has 4 problems.  Please write your code in the designated portion of the autograded answer cell.

For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](jupyter.png)

©2020 Luis Morales. All rights reserved. 

## Problem 1 (1 point)

Create a function that checks if the input number is prime or not. Function value will be True if argument is a prime number and False otherwise.

<!-- grade_id: is_prime, points: 0, grade: false, solution: true -->
```python
def is_prime(num):
    """check if the input number is prime or not.
    A positive integer greater than 1 which has no other factors except 1 and the number itself is called a prime number.
    in your solution, you can consider num is already verified as a positive integer!
    """
    
    if(type(num)!=int):
        raise ValueError('input is not an integer')
    if(num<=0):
        raise ValueError('input is not an integer')
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: is_prime_correct_input, points: 1, grade: true, solution: false -->
```python
"""Check that is_prime function returns the correct output for several inputs"""
assert is_prime(3) == True
assert is_prime(11) == True
assert is_prime(21) == False
assert is_prime(1) == False
```

## Problem 2 (1 point)

Write a function that helps identify a criminal(s) among a list of six suspects (designated A through F). The police has been conducting an investigation and they have reached the following conclusions: 

1. There is at least one criminal among A and B;
2. There are at least two criminals among A, E and F;
3. A and D can not be both criminals;
4. B and C may be both criminals or they are both innocent;
5. There is only one criminals in C and D;
6. If D is innocent, E is also innocent.

The output of your function must be a list of 6 integers with each digit corresponding to one of the suspects. If the \\(i^{th}\\) element of this returned list is 0 if the \\(i^{th}\\) suspect is not a criminal.  Otherwise, return a 1. 

<!-- grade_id: whoiscriminal, points: 0, grade: false, solution: true -->
```python
def whoIsCriminal():
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
whoIsCriminal()
```

<!-- grade_id: correct_cubics, points: 1, grade: true, solution: false -->
```python
assert whoIsCriminal() == [1, 1, 1, 0, 0, 1]
```

## Problem 3 (1 point)

Given a list of size n, find the majority element. The majority element is the element that appears more than n/2 times.

You may assume that the list is non-empty and the majority element always exist in the list.

<!-- grade_id: majorityelement, points: 0, grade: false, solution: true -->
```python
def majorityElement(nums):
   
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: checkmajorityelement, points: 1, grade: true, solution: false -->
```python
"""Check that majorityElement function returns the correct output"""
assert (majorityElement([3,7,3]) == 3)
assert (majorityElement([2,2,1,1,1,2,2]) == 2)
assert (majorityElement([3,2,3]) == 3)
assert (majorityElement(['abc','abc','dfdwfad']) == 'abc')

```

## Problem 4 (1 point)

Write a function that encodes an input string by spliting it by the middle and reversing the two substrings.

<!-- grade_id: encodestring, points: 0, grade: false, solution: true -->
```python
def encodeString(ori_str):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: encodestringtest, points: 1, grade: true, solution: false -->
```python
assert encodeString('AaAaBbBb') == 'aAaAbBbB'
assert encodeString('This line must be printed backwards and splitted in the middle.') \
    == 'wkcab detnirp eb tsum enil sihT.elddim eht ni dettilps dna sdra'
assert encodeString(encodeString('This lines must be printed backwards and splitted in the middle.')) \
    == 'This lines must be printed backwards and splitted in the middle.'
assert encodeString('This is a very secret message!') \
    == ' yrev a si sihT!egassem terces'
assert encodeString(encodeString('This is a very secret message!')) \
    == 'This is a very secret message!'
assert encodeString('Happy programming~') \
    == 'orp yppaH~gnimmarg'
assert encodeString(encodeString('Happy programming~')) \
    == 'Happy programming~'
```